# Week 3 Problem Set

## Cohort Sessions

**CS1.** You have implemented factorial using iteration in Problem Set 1. Now, implement the factorial problem using a recursion. The function should takes in an Integer input and returns and Integer output which is the factorial of the input. Recall that:

$n! = n \times (n-1) \times (n-2) \times \ldots \times 2 \times 1$

You should consider the case when $n$ is zero and one as well. 

```python
# recursion
def factorial_recursion(n):
    if n == 1 or n == 0:
        return 1
    else:
        return n * factorial_recursion(n-1)
```

<!-- grade_id: cs11, points: 1, grade: true, solution: false -->
```python
assert factorial_recursion(0) == 1
assert factorial_recursion(1) == 1
assert factorial_recursion(5) == 120
assert factorial_recursion(7) == 5040
assert factorial_recursion(11) == 39916800
```

<!-- grade_id: cs12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS2.** *Helper Function:* Write a function `palindrome(s)`to check if the string `s` is a [Palindrome](https://en.wikipedia.org/wiki/Palindrome). To do this, write another function `is_palindrome(s, left, right)` where `left` and `right` are indices from the left and from the right to check the character in `str`. Use recursion instead of iteration in this problem.


```python
def palindrome(s):
    return is_palindrome(s, 0, len(s)-1)

def is_palindrome(s, left, right):
    if right <=  left:
        return True
    elif s[left] != s[right]:
        return False
    else:
        return is_palindrome(s,left+1,right-1)
```

<!-- grade_id: cs21, points: 1, grade: true, solution: false -->
```python
assert not palindrome("moon") 
assert palindrome("noon") 
assert palindrome("a a") 
assert palindrome("ada") 
assert not palindrome("ad a") 
```

<!-- grade_id: cs22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS3.** *Towers of Hanoi:* Write a function `move_disks(n, from_tower, to_tower, aux_tower)` which returns an array of String for the movement of disks that solves the [Towers of Hanoi problem](https://en.wikipedia.org/wiki/Tower_of_Hanoi). 
* The first argument `n` is an Integer input that gives information on the number of disk. 
* The second argument `from_tower` is a String which is the label of the origin tower.
* The third argument `to_tower` is a String which is the label of the destination tower.
* The last argument `aux_tower` is a String which is the label of the auxilary tower.


```python
# def move_disks(n, from_tower, to_tower, aux_tower):
result = []
def move_disks(n, from_tower, to_tower, aux_tower):
    global result
    if n == 1:
        return [f'Move disk {n:} from {from_tower:} to {to_tower:}.']
    else:
        result = move_disks(n-1, from_tower, aux_tower, to_tower)
        result += [f'Move disk {n:} from {from_tower:} to {to_tower:}.']
        result += move_disks(n-1, aux_tower, to_tower, from_tower)
        return result
    
result = move_disks(2, "A", "B", "C")
print(result)
```

Output:
```
['Move disk 1 from A to C.', 'Move disk 2 from A to B.', 'Move disk 1 from C to B.']
```

<!-- grade_id: cs31, points: 1, grade: true, solution: false -->
```python
result = move_disks(3, "A", "B", "C")
print(result)
assert result == ["Move disk 1 from A to B.", "Move disk 2 from A to C.", "Move disk 1 from B to C.", "Move disk 3 from A to B.", "Move disk 1 from C to A.", "Move disk 2 from C to B.", "Move disk 1 from A to B."]

result = move_disks(4, "A", "B", "C")
print(result)
assert result == ["Move disk 1 from A to C.", "Move disk 2 from A to B.", "Move disk 1 from C to B.", "Move disk 3 from A to C.", "Move disk 1 from B to A.", "Move disk 2 from B to C.", "Move disk 1 from A to C.", "Move disk 4 from A to B.", "Move disk 1 from C to B.", "Move disk 2 from C to A.", "Move disk 1 from B to A.", "Move disk 3 from C to B.", "Move disk 1 from A to C.", "Move disk 2 from A to B.", "Move disk 1 from C to B."]

```

Output:
```
['Move disk 1 from A to B.', 'Move disk 2 from A to C.', 'Move disk 1 from B to C.', 'Move disk 3 from A to B.', 'Move disk 1 from C to A.', 'Move disk 2 from C to B.', 'Move disk 1 from A to B.']
['Move disk 1 from A to C.', 'Move disk 2 from A to B.', 'Move disk 1 from C to B.', 'Move disk 3 from A to C.', 'Move disk 1 from B to A.', 'Move disk 2 from B to C.', 'Move disk 1 from A to C.', 'Move disk 4 from A to B.', 'Move disk 1 from C to B.', 'Move disk 2 from C to A.', 'Move disk 1 from B to A.', 'Move disk 3 from C to B.', 'Move disk 1 from A to C.', 'Move disk 2 from A to B.', 'Move disk 1 from C to B.']
```

<!-- grade_id: cs32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS4.** *Merge Sort:* Write functions to implement the Merge Sort algorithm. The first function `mergesort(array)` should takes in an array of Integers in `array`. The function should sort the array in place. The second function `merge(array, p, q, r)` should implements the merge procedure. This function takes in an array of Integers in. `array`, the starting index for the left array `p`, the ending index for the left array `q`, and the ending index for the right array `r`. You can use a helper function for your recursion if needed.

```python
def merge(array, p, q, r):
    n1 = q - p + 1
    n2 = r - q
    left_array = array[p:q+1]
    right_array = array[q+1:r+1]
    
    i=0
    j=0
    k=p
    
    while (i < n1) and (j < n2):
        if left_array[i] <= right_array[j]:
            array[k] = left_array[i]
            i += 1
        else:
            array[k] = right_array[j]
            j +=1
        k += 1
    while (i < n1):
        array[k] = left_array[i]
        i += 1
        k += 1
    while (j < n2):
        array[k] = right_array[j]
        j += 1
        k += 1

def mergesort_recursive(array, p, r):
    if p < r:
        q = (p + r) // 2
        mergesort_recursive(array, p, q)
        mergesort_recursive(array, q+1, r)
        merge(array, p, q, r)
    
    
def mergesort(array):
    mergesort_recursive(array, 0, len(array)-1)
    print(array)


```

<!-- grade_id: cs41, points: 1, grade: true, solution: false -->
```python
input_array = [5, 2, 4, 7, 1, 3, 2, 6]
mergesort(input_array)
assert input_array == [1, 2, 2, 3, 4, 5, 6, 7]
```

Output:
```
[1, 2, 2, 3, 4, 5, 6, 7]
```

<!-- grade_id: cs42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```