# Week 12 Problem Set

## Homework

**HW1.** *Comments:* Write a state machine whose inputs are the characters of a string. The string contains the code for a computer program. The output of the state machine are either:
- the input character if it is part of a comment, or
- `None`, otherwise.

Comment starts with a `#` character and continue to the end of the current line. If you want to create a string that contains a new line character, you can use `\n`.

For example,
```
inpstr = "def func(x): # comment\n    return 1"
m = CommentsSM()
print(m.transduce(inpstr))
```

The expected output is:
```
[None, None, None, None, None, None, None, None, None, None, None, None, None, "#", " ", "c", "o", "m", "m", "e", "n", "t", None, None, None, None, None, None, None, None, None, None, None, None, None]
```

You should start by drawing a state transition diagram indicating the states and what inputs cause transition to which other states. Use the test case above to determine if your state transition diagram is correct. You should begin writing your program only when you are confident that your diagram is correct. 

```python
# Copy over the implementation of StateMachine from Cohort
from abc import ABC, abstractmethod

class StateMachine(ABC):
    def start(self):
        self.state = self.start_state

    def step(self, inp):
        ns, o = self.get_next_values(self.state, inp)
        self.state = ns
        return o
        
    def transduce(self, inp_list):
        output = []
        self.start()
        for inp in inp_list:
            if not self.is_done():
                o = self.step(inp)
                output.append(o)
        return output
        
    @abstractmethod
    def get_next_values(self, state, inp):
        pass

    def done(self, state):
        return False

    def is_done(self):
        return self.done(self.state)

```

```python
class CommentsSM(StateMachine):
    def __init__(self):
        self.start_state = "off"

    def get_next_values(self, state, inp):
        if state == "off":
            if inp == "#":
                next_state = "on"
                output = inp
            else:
                next_state = "off"
                output = None
        elif state == "on":
            if inp == "\n" :
                next_state = "off"
                output = None
            else:
                next_state = "on"
                output = inp
        return next_state, output
```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
inpstr = "def func(x): # comment\n    return 1"
m = CommentsSM()
out = m.transduce(inpstr)
assert out == [None, None, None, None, None, None, None, None, None, None, None, None, None, "#", " ", "c", "o", "m", "m", "e", "n", "t", None, None, None, None, None, None, None, None, None, None, None, None, None]
```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** *First Word:* Write a state machine whose inputs are the characters of a string and which outputs either:
- the input character if it is part of the first word on a line, or
- `None`, otherwise

For the purposes here, a word is any sequence of consecutive characters that does not contain spaces or end-of-line characters. In this problem, comments have no special status. This means that if the line begins with `# `, then the first word is `#`. 

For example, 
```
inpstr = "def func(x): # comment\n    return 1"
m = FirstWordSM()
print(m.transduce( inpstr))
```

The expected output is:
```
["d", "e", "f", None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, "r", "e", "t", "u", "r", "n", None, None]
```

```python
class FirstWordSM(StateMachine):
    def __init__(self):
        self.start_state = 0
        
    def get_next_values(self, state, inp):
        if state == 0:
            if inp == "\n" or inp == " ":
                output = None
                next_state = 0
            else:
                output = inp
                next_state = 1
        elif state == 1:
            if inp == " ":
                output = None
                next_state = 2
            elif inp == '\n':
                output = None
                next_state = 0
            else:
                output = inp
                next_state = 1
        elif state == 2:
            if inp == "\n":
                output = None
                next_state = 0
            else:
                output = None
                next_state = 2
        return next_state, output
```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
inpstr = "def func(x): # comment\n    return 1"
m = FirstWordSM()
out = m.transduce(inpstr)
assert out == ["d", "e", "f", None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, None, "r", "e", "t", "u", "r", "n", None, None]

```

<!-- grade_id: hw22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW3.** *Robot:* Write a State Machine class that represent a robot. The dimension of the world and the robot initial position should be specified during the class instantiation. The robot can take in the following input:
- "left"
- "right"
- "up"
- "down"

The initial position of the robot is specified during the object instantiation and the input should modify the position of the robot. The robot position must not change if it exceed the boundary. At each step, the robot should output the updated position. 

```python
class Position:
  
    def __init__(self, x=0, y=0):
        self.x = x
        self.y = y
  
    def __str__(self):
        return f"({self.x:}, {self.y:})"

class Dimension:

    def __init__(self, width=0, height=0):
        self.width = width
        self.height = height
  
    def __str__(self):
        return f"width: {self.width:}, height: {self.height:}"
```

```python
class RobotSM(StateMachine):

    def __init__(self, init_pos, dimension):
        self.world_dim = dimension
        self.start_state = init_pos
        self.start()
  
    def get_next_values(self, state, inp):
        if inp == 'left' and state.x != 0:
            next_state = Position(state.x - 1, state.y)
        elif inp == 'up' and state.y != self.world_dim.height:
            next_state = Position(state.x, state.y + 1)
        elif inp == 'down' and state.y != 0:
            next_state = Position(state.x, state.y - 1)
        elif inp == 'right' and state.y != self.world_dim.width:
            next_state = Position(state.x + 1, state.y)
        return next_state, None
```

<!-- grade_id: hw31, points: 1, grade: true, solution: false -->
```python
robot = RobotSM(Position(0, 0), Dimension(5, 5))
robot.start()
robot.transduce(["right", "right", "up", "up", "up", "left", "down"])
pos = robot.state
assert pos.x == 1 and pos.y == 2
```

Output:
```
<ipython-input-3-5de78d32a267> in step(self, inp)
      7 
      8     def step(self, inp):
----> 9         ns, o = self.get_next_values(self.state, inp)
     10         self.state = ns
     11         return o

<ipython-input-11-9f957d04c186> in get_next_values(self, state, inp)
     15         elif inp == 'right' and state.y != self.world_dim.width:
     16             next_state = Position(state.x + 1, state.y)
---> 17         return next_state, output

NameError: name 'output' is not defined
```

<!-- grade_id: hw32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW4.** *Search SM:* Write a function `sm_search` that takes in the following arguments:
- `sm_to_search`: is the State Machine instance to search. This argument is of the type `MapSM` as defined in CS4. You should use the `get_next_values()` of this State Machine instance to explore the next state in your search.
- `initial_state`: is the start state of the search. If it is not provided, it should be assigned to the `start_state` of sm_search. 
- `goal_test`: is a function that returns `True` if the argument is the end state of the search. If it is not provided, it should be eassigned to the `done` function of the state machine.

This function performs a **breadth-first-search** algorithm to explore the next states. 

The output is a `list` of `Step` instances from the `init_state` to the end state which is determined by the `goal_test` function.

This problem requires you to complete the following:
- `Queue` class from Problem Set 4 HW2.
- `MapSM` class in CS4.
- `SearchNode` and `Step` classes in CS5.

```python
# Copy over the implementations of Queue from PS4 HW2
class Queue:
    def __init__(self):
        self.left_stack = Stack()
        self.right_stack = Stack()
        
    def enqueue(self,item):
        self.right_stack.push(item)
        
        
    def dequeue(self):
        if self.left_stack.is_empty:
            while not self.right_stack.is_empty:
                item = self.right_stack.pop()
                self.left_stack.push(item)
                
        return self.left_stack.pop()
        
        
    def peek(self):
        if self.left_stack.is_empty:
            if not self.right_stack.is_empty:
                return self.right_stack.items[0]
        else:
            return self.left_stack.peek()
        
    @property
    def is_empty(self):
        return self.left_stack.is_empty and self.right_stack.is_empty
    
    @property
    def size(self):
        return self.left_stack.size + self.right_stack.size
```

```python
# Copy over the implementation of StateSpaceSearch from Cohort
from abc import abstractmethod

class StateSpaceSearch(StateMachine):
    @property
    @abstractmethod
    def statemap(self):
        pass

    @property
    @abstractmethod
    def legal_inputs(self):
        pass

```

```python
# Copy over the implementation of MapSM from Cohort
class MapSM(StateSpaceSearch):
    
    def __init__(self, start):
        self.start_state = start
        
    @property
    def statemap(self):
        statemap = {"S": ["A", "B"],
                    "A": ["S", "C", "D"],
                    "B": ["S", "D", "E"],
                    "C": ["A", "F"],
                    "D": ["A", "B", "F", "H"],
                    "E": ["B", "H"],
                    "F": ["C", "D", "G"],
                    "H": ["D", "E", "G"],
                    "G": ["F", "H"]}
        return statemap

    @property
    def legal_inputs(self):
        max_neighbour = -1
        for state, neighbours in self.statemap.items():
            if max_neighbour < len(neighbours):
                max_neighbour = len(neighbours)
        
        # max_neighbour is just an integer
        # we need to return a set 
        return set(range(0, max_neighbour))

    # eg if state == S, and inp == 1, 
    # return "B", "B"
    def get_next_values(self, state, inp):
        # use a method that will not result in an error if the key isn't found
        neighbours = self.statemap.get(state, None)
        # why not: neighbours = self.statemap[state]?
        # because if state is not a valid key, the above will return error

        # default values
        next_state = state
        output = state 

        # if neighbours is None:
        #     return next_state, output
        
        if neighbours != None and inp < len(neighbours):
            next_state = neighbours[inp]
            output = next_state
        
        return next_state, output

```

```python
# Copy over the implementations of Step and SearchNode from Cohort
class Step:
    def __init__(self, action, state):
        self.action = action
        self.state = state
    
    def __eq__(self, other):
        return self.action == other.action and self.state == other.state
  
    def __str__(self):
        return f"action: {self.action:}, state: {self.state:}"

class SearchNode:
    def __init__(self, action, state, parent):
        self.action = action
        self.state = state
        self.parent = parent
  
    # @return
    # list of Step instances
    # Example: S -> A -> C
    # C.path() SHOULD return [Step(None, S), Step(ActionA, A), Step(ActionC, C)]
    # C.path() go to else clause:
        # return A.path() + [Step(ActionC, C)]
    # Requires us to compute A.path() before C.path() can return
    # A.path() go to else clause:
        # return S.path() + [Step(ActionA, A)]
    # Requires us to compute S.path() before A.path() can return
    # S.path() go to if clause:
        # return [Step(None, S)] to A.path()
    # A.path() now executes:
        # return [Step(None, S)] + [Step(ActionA, A)] to C.path()
    # C.path() now executes:
        # return [Step(None, S)] + [Step(ActionA, A)] + [Step(ActionC, C)] 

    def path(self):
        # base state
        if self.parent is None:
            return [Step(self.action, self.state)]
        else:
            return self.parent.path() + [Step(self.action, self.state)]
        
    # @args
    # state --> string, or anything that represents a state
    # @return
    # boolean
    def in_path(self, state):
        if self.state == state:
            return True
        elif self.parent == None:
            return False 
        else:
            # recursion
            # just like a linked list,
            # check if the asked state exists
            # in your ancestor line
            return self.parent.in_path(state)
  
    def __eq__(self, other):
        if self is None and other is None:
            return True
        elif self is None:
            return False
        elif other is None:
            return False
        else:
            return self.state == other.state and self.parent == other.parent and \
                   self.action == other.action

```

```python
def sm_search(sm_to_search, initial_state=None, goal_test=None):
    # check if initial_state is provided
    # if it is, use it
    # otherwise, get the start state of sm_to_search
    if initial_state == None:
        # replace None to take the start state of sm_to_search
        init_state = None
        ###
        ### YOUR CODE HERE
        ###
    else:
        init_state = initial_state
  
    # check if goal_test is provided
    # if it is, use it
    # otherwise, use the done method as the goal function
    # taken from sm_to_search
    if goal_test == None:
        goal_func = sm_to_search.done
    else:
        goal_func = goal_test
  
    # create a Queue instance to store the node to explore
    # replace the None below
    agenda = None
    ###
    ### YOUR CODE HERE
    ###
  
    # if the initial state is the goal state, 
    # then we are done and exit
    if goal_func(init_state):
        return [Step(None, init_state)]
  
    # otherwise, add the current node into the agenda 
    agenda.enqueue(SearchNode(None, init_state, None))
    
    # explore as long as the Queue is not empty
    while not agenda.is_empty:
        
        # replace None to take out the parent from the Queue
        parent = None
        ###
        ### YOUR CODE HERE
        ###
        
        # create a list to keep track which child state have been explored
        new_child_state = []
        
        # get all the legal input values
        actions = sm_to_search.legal_inputs
        
        #iterate over all legal inputs
        for a in actions:
            # get the next possible state using the current action
            # call get_next_values to get the next state
            # replace the None below
            new_s = None
            ###
            ### YOUR CODE HERE
            ###
            
            # create a new search node from the new_s
            # replace the None below
            new_n = None
            ###
            ### YOUR CODE HERE
            ###
            
            # if the new state is the goal state, then we exit and return the path
            if goal_func(new_s):
                return new_n.path()
            
            # if the new state is already in the list of new child state, ignore it
            elif new_s in new_child_state:
                pass
            
            # if the new state is in the path of the current node, ignore it
            elif parent.in_path(new_s):
                pass
            
            # otherwise, add the new state into the list
            # and the new node into the Queue
            else:
                # step 1. add the new state into the new_child_state
                ###
                ### YOUR CODE HERE
                ###
                
                # step 2. add the new node into the Queue
                ###
                ### YOUR CODE HERE
                ###
                pass
    return None
```

<!-- grade_id: hw41, points: 1, grade: true, solution: false -->
```python
mapSM = MapSM("S")
ans = sm_search(mapSM , "S" , lambda s: s=="H" )
steps = [(step.action, step.state) for step in ans]
assert steps == [(None, "S"), (0, "A"), (2, "D"), (3, "H")]
for step in ans:
    print(step)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-5-61237f5b09a0> in <module>
----> 1 mapSM = MapSM("S")
      2 ans = sm_search(mapSM , "S" , lambda s: s=="H" )
      3 steps = [(step.action, step.state) for step in ans]
      4 assert steps == [(None, "S"), (0, "A"), (2, "D"), (3, "H")]
      5 for step in ans:

NameError: name 'MapSM' is not defined
```

<!-- grade_id: hw42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```