# Week 5 Problem Set

## Homework

**HW1.** *Dictionary:* Write two functions:
1. `count_degrees(G)`: which sums up the degrees of all vertices in the graph. The degree of a vertex is defined as the number of edges connected to a Vertex. The argument `G` is a dictionary that represents the graph.
2. `count_edges(G)`: which counts the number of edges in the graph. An edge is defined as a connection between two vertices. The argument `G` is a dictionary.

```python
def count_degrees(G):
    degrees = 0
    for item in G:
        degrees += len(G.get(item))
    return degrees

def count_edges(G):
    edge_tracker = []
    edge = ""
    for key in G:
        for Edges in G.get(key):
            edge = sorted(key + Edges)
            if edge not in edge_tracker:
                edge_tracker.append(edge)
    return len(edge_tracker)

```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
g1 = {"A": ["B", "E"],
      "B": ["A", "C"],
      "C": ["B", "D", "E"],
      "D": ["C"],
      "E": ["A", "C"]}

d = count_degrees(g1)
e = count_edges(g1)

assert d == 10
assert e == 5
```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** Create a class called `GraphSearch` which is a subclass of the class `Graph`. This class should override the method `_create_vertex(id)` and instantiate a `VertexSearch` object instead of `Vertex`.

```python
# Copy Vertex class from Cohort
class Vertex:
    def __init__(self, id=""):
        self.id = id
        self.neighbours = {}
        
    def add_neighbour(self, numb_vertex, weight = 0):
        self.neighbours[numb_vertex] = weight
    
    def get_neighbours(self):
        return list(self.neighbours.keys())
    
    def get_weight(self, neighbour):
        return self.neighbours.get(neighbour, None)
    
    def __eq__(self, other):
        return self.id == other.id
    
    def __lt__(self, other):
        return self.id < other.id
    
    def __hash__(self):
        return hash(self.id)
    
    def __str__(self):
        return f"Vertex {self.id} is connected to: "             + ', '.join((j.id) for j in self.get_neighbours())


```

```python
# Copy VertexSearch class from Cohort
import sys

class VertexSearch(Vertex):
    def __init__(self, id=""):
        super().__init__(id)
        self.colour = "white"
        self.d = sys.maxsize
        self.f = sys.maxsize
        self.parent = None


```

```python
#Copy Graph over from Cohort
class Graph:
    def __init__(self):
        self.vertices = {}
        
    def _create_vertex(self, id):
        return Vertex(id)
    
    def add_vertex(self, id):
        self.vertices[id] = self._create_vertex(id)
        pass
    
    def get_vertex(self, id):
        return self.vertices.get(id, None)
    
    def add_edge(self, start_v, end_v, weight = 0):
        if start_v not in self:
            vs = self.add_vertex(start_v)
        else:
            vs = self.vertices[start_v]
        
        if end_v not in self:
            ve = self.add_vertex(end_v)
        else:
            ve = self.vertices[end_v]
            
        vs.add_neighbour(ve, weight)
        pass
    
    def get_neighbours(self, id):
        v = self.get_vertex(id)
        if v is not None:
            result = []
            
            for i in v.get_neighbours():
                result.append(i.id)
            return result
        else:
            return None
    
    def __contains__(self, val):
        return val in self.vertices.keys()
    
    def __iter__(self):
        for k,v in self.vertices.items():
            yield v 
    @property
    def num_vertices(self):
        return len(self.vertices)

```

```python
class GraphSearch(Graph):
    ##BEGIN SOLUTION
    def _create_vertex(self, id):
        return VertexSearch(id)
    ##END SOLUTION
    pass
```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
g2 = GraphSearch()
g2.add_vertex("Z")
assert(type(g2.vertices["Z"]) == type(VertexSearch()))
```

<!-- grade_id: hw22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW3.** *Undirected Graph:* **You need to complete CS5 and CS6 before you do this homework.** Paste your answer for `Search2D` and `SearchBFS` classes into the cell below.

Create a subclass of `GraphSearch` class called `UGraphSearch` for undirected graphs. All edges in undirected graphs are bidirectional (i.e. vertex1 <-> vertex2). 
This means that you need to override the `add_edge()` method.

```python
#Copy Queue over from Wk4 Homework
class Queue:
    def __init__(self):
        self.__items = []
    
    def enqueue(self, item):
        self.__items.append(item)
    
    def dequeue(self):
        return self.__items.pop(0)
    
    def peek(self):
        return self.__items[0]
    
    @property
    def is_empty(self):
        if self.__items == []:
            return True
        return False
    
    @property
    def size(self):
        return len(self.__items)

```

```python
# copy Search2D over from CS5
import sys

class Search2D:
    def __init__(self, g):
        self.graph = g
    
    def clear_vertices(self):
        for v in self.graph.vertices.values():
            v.colour = "white"
            v.d = sys.maxsize
            v.f = sys.maxsize
            v.parent = None
            
    def __iter__(self):
        return iter([v for v in self.graph])
    
    def __len__(self):
        return len([v for v in self.graph.vertices])

```

```python
# copy SearchBFS over from CS6

class SearchBFS(Search2D):
    def search_from(self, start):
        self.clear_vertices()
        startv = self.graph.get_vertex(start)
        if startv is None:
            return 
        startv.colour = "grey" 
        startv.d = 0
        startv.parent = None
        q = Queue()
        q.enqueue(startv)
    
        while not q.is_empty:
            u = q.dequeue()
            for v in u.get_neighbours():
                if v.colour == "white":
                    v.colour = 'grey'
                    v.d = u.d + 1
                    v.parent = u
                    q.enqueue(v)
            u.colour = 'black'

    def get_shortest_path(self, start, dest):
        result = []
        self.get_path(start, dest, result)
        if result == []:
            return None
        return result
    
    def get_path(self, start, dest, result):
        u = self.graph.get_vertex(start)
        if u is None:
            return
        v = self.graph.get_vertex(dest)
        if v is None:
            return

        if u.d != 0:
            self.search_from(start)
        
        if u == v:
            result.append(start)
        elif v.parent is None:
            result.append("No Path")
        else:
            self.get_path(start, v.parent.id, result)
            result.append(dest)

```

```python
class UGraphSearch(GraphSearch):
    def add_edge(self, start_v, end_v, weight = 0):
        if start_v not in self:
            vs = self.add_vertex(start_v)
            ve = self.add_vertex(end_v)
        else:
            vs = self.vertices[start_v]
            ve = self.vertices[end_v]
        
        if end_v not in self:
            ve = self.add_vertex(end_v)
            vs = self.add_vertex(start_v)
        else:
            ve = self.vertices[end_v]
            vs = self.vertices[start_v]
        vs.add_neighbour(ve, weight)
        ve.add_neighbour(vs, weight)   


```

<!-- grade_id: hw31, points: 1, grade: true, solution: false -->
```python
g2 = UGraphSearch()
assert g2.vertices == {} and g2.num_vertices == 0
g2.add_vertex("L")
g2.add_vertex("I")
g2.add_vertex("S")
g2.add_vertex("P")
assert g2.num_vertices == 4
assert "L" in g2
assert "I" in g2
assert "S" in g2
assert "P" in g2
g2.add_edge("L", "I")
g2.add_edge("I", "S")
g2.add_edge("S", "P")
assert sorted(g2.get_neighbours("L")) == ["I"]
assert sorted(g2.get_neighbours("I")) == ["L", "S"]
assert sorted(g2.get_neighbours("S")) == ["I", "P"]
assert sorted(g2.get_neighbours("P")) == ["S"]
```

<!-- grade_id: hw32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: hw33, points: 1, grade: true, solution: false -->
```python
g2 = UGraphSearch()
g2.add_vertex("r")
g2.add_vertex("s")
g2.add_vertex("t")
g2.add_vertex("u")
g2.add_vertex("v")
g2.add_vertex("w")
g2.add_vertex("x")
g2.add_vertex("y")
g2.add_vertex("z")
g2.add_edge("r", "s")
g2.add_edge("r", "v")
g2.add_edge("s", "w")
g2.add_edge("t", "u")
g2.add_edge("t", "x")
g2.add_edge("t", "w")
g2.add_edge("u", "t")
g2.add_edge("u", "x")
g2.add_edge("u", "y")
g2.add_edge("v", "r")
g2.add_edge("w", "s")
g2.add_edge("w", "t")
g2.add_edge("w", "x")
g2.add_edge("x", "w")
g2.add_edge("x", "t")
g2.add_edge("x", "u")
g2.add_edge("x", "y")
g2.add_edge("y", "u")
g2.add_edge("y", "x")
gs = SearchBFS(g2)
gs.search_from("s")
assert gs.graph.get_vertex("s").d == 0
assert gs.graph.get_vertex("t").d == 2
assert gs.graph.get_vertex("u").d == 3
assert gs.graph.get_vertex("v").d == 2
assert gs.graph.get_vertex("w").d == 1
assert gs.graph.get_vertex("x").d == 2
assert gs.graph.get_vertex("y").d == 3
assert gs.graph.get_vertex("r").d == 1
ans = gs.get_shortest_path("s", "u")
assert ans == ["s", "w", "t", "u"] or ans == ["s", "w", "x", "u"]
ans = gs.get_shortest_path("v", "s")
assert ans == ["v", "r", "s"]
ans = gs.get_shortest_path("v", "y")
assert ans == ["v", "r", "s", "w", "x", "y"]
```

<!-- grade_id: hw34, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW4.** *Depth-First-Search:* Create a class `SearchDFS` as a child class of `Search2D` to implement Depth-First-Search algorithm. You should add one additional attribute:
- `time`: which is an attribute to record the discovery time and the finishing time.

The class should also implement the following methods:
- `search()`: which modifies the vertices' properties such as `colour`, `d`, and `parent` following Depth-First-Search algorithm.
- `dfs_visit(vert)`: which is the recursive method for visiting a vertex in Depth-First-Search algorithm.

```python
import sys

class SearchDFS(Search2D):
    def __init__(self, g):
        self.graph = g
        self.time = 0
      
    def search(self):
        self.clear_vertices()
        self.time = 0
        for vertex in self.graph.vertices.values():
            if vertex.colour == "white":
                self.dfs_visit(vertex) 

    def dfs_visit(self, vert):
        self.time += 1
        vert.d = self.time
        vert.colour = "grey"
        for neighbours in vert.get_neighbours():
            if neighbours.colour == "white":
                neighbours.parent = vert
                self.dfs_visit(neighbours)
        vert.colour = "black"
        self.time += 1
        vert.f = self.time

```

<!-- grade_id: hw41, points: 1, grade: true, solution: false -->
```python
g4 = GraphSearch()
g4.add_vertex("e")
g4.add_vertex("m")
g4.add_vertex("a")
g4.add_vertex("c")
g4.add_vertex("s")
g4.add_edge("e", "m")
g4.add_edge("m", "a")
g4.add_edge("a", "c")
g4.add_edge("c", "s")
gs = SearchDFS(g4)
gs.search()
assert gs.graph.get_vertex("e").parent == None 
assert gs.graph.get_vertex("m").parent == gs.graph.get_vertex("e")

assert gs.graph.get_vertex("m").d == 2 and gs.graph.get_vertex("a").f == 8
assert gs.graph.get_vertex("c").d == 4 and gs.graph.get_vertex("s").f == 6
```

<!-- grade_id: hw42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW5.** *Topological Sort:* Create a function that takes in a `SearchDFS` object to perform a topological sort:
- `topological_sort(g)`: which takes in a `SearchDFS` object and returns a list of `VertexSearch` objects sorted based on their `f` property. This method should call the `search()` method of the `SearchDFS` to first calculate the `f` property of the vertices. Note that you need to copy the `SearchDFS` object of your input argument so as not to mutate the object.

```python
import sys
import copy

def topological_sort(g):
    gs = g
    gs.search()
    sort = []

    for vertex in gs.graph.vertices.values():
        sort.append(vertex) 
    swap = True
    arLength = len(sort)
    temp = None
    while (swap == True):
        swap = False
        for i in range (0,arLength-1):
            if sort[i].f < sort[i + 1].f:
                temp = sort[i]
                sort[i] = sort[i + 1]
                sort[i + 1] = temp
                swap = True   
    return sort 
```

<!-- grade_id: hw51, points: 1, grade: true, solution: false -->
```python
import copy
g = GraphSearch()
g.add_vertex("3/4 cup milk")
g.add_vertex("1 egg")
g.add_vertex("1 tbl oil")
g.add_vertex("1 cup mix")
g.add_vertex("heat syrup")
g.add_vertex("heat griddle")
g.add_vertex("pour 1/4 cup")
g.add_vertex("turn when bubbly")
g.add_vertex("eat")
g.add_edge("3/4 cup milk", "1 cup mix")
g.add_edge("1 egg", "1 cup mix")
g.add_edge("1 tbl oil", "1 cup mix")
g.add_edge("1 cup mix", "heat syrup")
g.add_edge("1 cup mix", "pour 1/4 cup")
g.add_edge("heat griddle", "pour 1/4 cup")
g.add_edge("pour 1/4 cup", "turn when bubbly")
g.add_edge("turn when bubbly", "eat")
g.add_edge("heat syrup", "eat")
gs = SearchDFS(g)  

path = topological_sort(gs)
ans = [v.f for v in copy.copy(path)]
assert ans == [18, 16, 14, 12, 11, 10, 9, 6, 5]
ans = [v.id for v in copy.copy(path)]
assert ans == ['heat griddle', '1 tbl oil', '1 egg', '3/4 cup milk', '1 cup mix', 'pour 1/4 cup', 'turn when bubbly', 'heat syrup', 'eat']
```

<!-- grade_id: hw52, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```