# Week 6 Problem Set

## Homeworks

**HW1.** Extend the class `Fraction` to implement the other operators: `- * < <= > >=`.

```python
def gcd(a, b):
    if b == 0:
        return a
    else:
        return gcd(b, a % b)


class Fraction:
    def __init__(self, num, den):
        ###BEGIN SOLUTION
        self.num = num # set attributes via the SETTER
        self.den = den
        ###END SOLUTION
        pass
    
    @property
    def num(self):
        ###BEGIN SOLUTION
        return self._numattr
        ###END SOLUTION
        pass
    
    @num.setter
    def num(self, val):
        ###BEGIN SOLUTION 
        self._numattr = int(val) # this creates attribute _num IF it does not exist already, else this modifies existing attr _num
        ###END SOLUTION
        pass
    
    @property
    def den(self):
        ###BEGIN SOLUTION
        return self._denattr
        ###END SOLUTION
        pass
    
    @den.setter
    def den(self, val):
        ###BEGIN SOLUTION
        if (int(val) == 0):
            self._denattr = 1
        else:
            self._denattr = int(val)
        ###END SOLUTION
        pass
    
    def __str__(self):
        ###BEGIN SOLUTION
        # 3/5
        return f"{self.num}/{self.den}"
        ###END SOLUTION
        pass
    
    def simplify(self):
        common = gcd(self.num, self.den)
        num = self.num//common
        den = self.den//common
        return Fraction(num, den)

    def __sub__(self, other):
        a = self.num
        b = self.den
        c = other.num
        d = other.den
        num = a*d - b*c
        den = b*d
        return Fraction(num, den).simplify()
    
    def __add__(self, other):
        a = self.num
        b = self.den
        c = other.num
        d = other.den
        num = a*d + b*c
        den = b*d
        return Fraction(num,den).simplify()

    
    def __mul__(self, other):
        res_num = self.num * other.num
        res_den = self.den * other.den
        return Fraction(res_num, res_den).simplify()
        pass
    
    def __eq__(self, other):
        left = self.simplify()
        right = other.simplify()
        return left.num == right.num and left.den == right.den
        pass
    
    def __lt__(self, other):
        gcd_val = gcd(self.den, other.den)
        if self.den > other.den:
            other.den *= gcd_val
            other.num *= gcd_val
        elif self.den < other.den:
            self.den *= gcd_val
            self.num *= gcd_val
        if self.num < other.num:
            return True
        pass
    
    def __le__(self, other):
        gcd_val = gcd(self.den, other.den)
        if self.den > other.den:
            other.den *= gcd_val
            other.num *= gcd_val
        elif self.den < other.den:
            self.den *= gcd_val
            self.num *= gcd_val
        if self.num <= other.num:
            return True
        pass
    
    def __gt__(self, other):
        gcd_val = gcd(self.den, other.den)
        if self.den > other.den:
            other.den *= gcd_val
            other.num *= gcd_val
        elif self.den < other.den:
            self.den *= gcd_val
            self.num *=gcd_val
            
        if self.num > other.num:
            return True


    def __ge__(self, other):
        gcd_val = gcd(self.den, other.den)
        if self.den > other.den:
            other.den *= gcd_val
            other.num *= gcd_val
        elif self.den < other.den:
            self.den *= gcd_val
            self.num *=gcd_val
        if self.num >= other.num:
            return True
    

```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
f1 = Fraction(3, 4)
f2 = Fraction(1, 2)
f3 = f1 - f2
assert f3 == Fraction(1, 4)
f4 = f1 * f2
assert f4 == Fraction(3, 8)
assert f2 < f1
assert f2 <= f2
assert f1 > f3
assert f3 >= f3
```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
# Copy your solution from the Cohort problem
class MixedFraction(Fraction):
    def __init__(self, top, bot, whole=0):
        num = whole * bot + top
        super().__init__(num, bot)
    
    def get_three_numbers(self):
        whole = self.num // self.den
        top = self.num % self.den
        bot = self.den
        return top, bot, whole

    def __str__(self):
        (top, bot, whole) = self.get_three_numbers()
        if whole > 0:
            return (f"{whole} {top}/{bot}")
        else:
            return super().__str__()

```

<!-- grade_id: hw13, points: 1, grade: true, solution: false -->
```python
mf1 = MixedFraction(5, 3)
assert mf1.num == 5 and mf1.den == 3
assert mf1.get_three_numbers() == (2, 3, 1)
mf2 = MixedFraction(2, 3, 1)
assert mf2.num == 5 and mf2.den == 3

result = mf1 + mf2
assert result.num == 10 and result.den == 3

result = mf1 * mf2
assert result.num == 25 and result.den == 9

mf3 = MixedFraction(1, 2, 1)
result = mf1 - mf3
assert result.num == 1 and result.den == 6

assert str(mf1) == "1 2/3"
```

<!-- grade_id: hw14, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** Write a class called `EvaluateFraction` that evaluates postfix notation implemented using Dequeue data structures only. Postfix notation is a way of writing expressions without using parenthesis. For example, the expression `(1+2)*3` would be written as `1 2 + 3 *`. The class `EvaluateFraction` has the following method:
- `input(inp)`: which pushes the input input one at a time. For example, to create a postfix notation `1 2 + 3 *`, we can call this method repetitively, e.g. `e.input('1'); e.input('2'); e.input('+'); e.input('3'); e.input('*')`. Notice that the input is of String data type. 
- `evaluate()`: which returns the output of the expression.
- `get_fraction(inp)`: which takes in an input string and returns a `Fraction` object. 

Postfix notation is evaluated using a Stack. Since `Dequeue` can be used for both Stack and Queue, we will implement using `Dequeue`. The input streams from `input()` are stored in a Queue, which we will again implement using Dequeue. If the output of the Queue is a number, the item is pushed into the stack. If it is an operator, we will apply the operator to the two top most item n the stacks and push the result back into the stack. 

```python
class Stack:
    def __init__(self):
        self.items = []
        
    @property
    def is_empty(self):
        return self.items == []
    
    @property
    def size(self):
        return len(self.items)
    
    def push(self, items):
        self.items.append(items)
        
    def pop(self):
        return self.items.pop()
    
    def peek(self):
        return self.items[-1]

```

```python
class Queue:
    def __init__(self):
        self.left_stack = Stack()
        self.right_stack = Stack()
    @property
    def is_empty(self):
        return self.left_stack.is_empty and self.right_stack.is_empty
    @property
    def size(self):
        return self.left_stack.size + self.right_stack.size
    def enqueue(self, item):
        self.right_stack.push(item)
    def dequeue(self):
        if self.left_stack.is_empty:
            self.left_stack.items = self.right_stack.items[:]
            self.left_stack.items.reverse()
            self.right_stack.items = []
        return self.left_stack.pop()
    def peek(self):
        if self.left_stack.is_empty:
            if not self.right_stack.is_empty:
                return self.right_stack.items[0]
        else:
            return self.left_stack.peek()
```

```python
class Dequeue(Queue):
    def add_front(self, item):
        self.left_stack.push(item)
    def remove_front(self):
        return self.dequeue()
    def add_rear(self, item):
        self.enqueue(item)
    def remove_rear(self):
        if self.right_stack.is_empty:
            self.left_to_right()
            return self.right_stack.pop()
    def left_to_right(self):
        while not self.left_stack.is_empty:
            item = self.left_stack.pop()
            self.right_stack.push(item)
    def peek_front(self):
        return self.peek()
    def peek_rear(self):
        if self.right_stack.is_empty:
            self.left_to_right()
        return self.right_stack.peek()
```

```python
class EvaluateFraction:

    operands = "0123456789"
    operators = "+-*/"
    
    def __init__(self):
        self.expression = Dequeue()
        self.stack = Dequeue()
    
    def input(self, item):
        self.expression.add_rear(item)
    def evaluate(self):
        #evaluate(): which returns the output of the expression.
        while not self.expression.is_empty:
            i = self.expression.remove_front()
            # print(f"i = {i}")
            if i in self.operands:
                # print("f if i in self.operands ran, with i = {i}")
                self.stack.add_front(i)
            elif "/" in i:
                # print(f"/ in i ran with i = {i}")
                self.stack.add_front(self.get_fraction(i))
                # print(f"stack.add_front added with {self.stack.peek_front()}")
            elif i in self.operators:
                first_val = self.stack.remove_front()
                print(f"first_val = {first_val}")
                operator = i
                # print(f"operator = {i}")
                sec_val = self.stack.remove_front()
                print(f"sec_val = {sec_val}")
                res = self.process_operator(first_val, sec_val, operator)
                print(f"res = {res}")
                self.stack.add_front(res)
        return self.stack.remove_rear()
    
    def get_fraction(self, inp):
        i = inp.split('/')
        return Fraction(i[0], i[1])
    
    def process_operator(self, op1, op2, op):
        if op == '+':
            return op1 + op2
        elif op == '-':
            return op2 - op1
        elif op == '*':
            return op1 * op2
# pe = EvaluateFraction()
# pe.input("1/2")
# pe.input("2/3")
# pe.input("+")
# assert pe.evaluate()==Fraction(7, 6)

```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
pe = EvaluateFraction()
pe.input("1/2")
pe.input("2/3")
pe.input("+")
assert pe.evaluate()==Fraction(7, 6)

pe.input("1/2")
pe.input("2/3")
pe.input("+")
pe.input("1/6")
pe.input("-")
assert pe.evaluate()==Fraction(1, 1)

pe.input("1/2")
pe.input("2/3")
pe.input("+")
pe.input("1/6")
pe.input("-")
pe.input("3/4")
pe.input("*")
assert pe.evaluate()==Fraction(3, 4)
```

Output:
```
first_val = 2/3
sec_val = 1/2
res = 7/6
first_val = 2/3
sec_val = 1/2
res = 7/6
first_val = 1/6
sec_val = 7/6
res = 1/1
first_val = 2/3
sec_val = 1/2
res = 7/6
first_val = 1/6
sec_val = 7/6
res = 1/1
first_val = 3/4
sec_val = 1/1
res = 3/4
```

<!-- grade_id: hw22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW3.** Modify HW2 so that it can work with MixedFraction. Write a class called `EvaluateMixedFraction` as a subclass of `EvaluateFraction`. You need to override the following methods:
- `get_fraction(inp)`: This function should be able to handle string input for MixedFraction such as `1 1/2` or `3/2`. It should return a `MixedFraction` object.
- `evaluate()`: This function should return `MixedFraction` object rather than `Fraction` object. 

```python
class EvaluateMixedFraction(EvaluateFraction):
    def get_fraction(self, inp):
        if ' ' in inp:
            i = inp.split (' ')
            whole_num = i[0]
            x = i[1].split('/')
            num = x[0]
            den = x[1]
            mixed_num = int(num) + int(den) * int(whole_num)
        else:
            z = inp.split('/')
            mixed_num = z[0]
            den = z[1]
        return Fraction(mixed_num, den)
    
    def evaluate(self):
        #evaluate(): which returns the output of the expression.
        while not self.expression.is_empty:
            i = self.expression.remove_front()
            # print(f"i = {i}")
            if i in self.operands:
                # print("f if i in self.operands ran, with i = {i}")
                self.stack.add_front(i)
            elif "/" in i:
                # print(f"/ in i ran with i = {i}")
                self.stack.add_front(self.get_fraction(i))
                # print(f"stack.add_front added with {self.stack.peek_front()}")
            elif i in self.operators:
                first_val = self.stack.remove_front()
                # print(f"first_val = {first_val}")
                operator = i
                # print(f"operator = {i}")
                sec_val = self.stack.remove_front()
                # print(f"sec_val = {sec_val}")
                res = self.process_operator(first_val, sec_val, operator)
                # print(f"res = {res}")
                self.stack.add_front(res)
        answer = self.stack.remove_rear()
        return MixedFraction(answer.num, answer.den)
```

<!-- grade_id: hw31, points: 1, grade: true, solution: false -->
```python
pe = EvaluateMixedFraction()
pe.input("3/2")
pe.input("1 2/3")
pe.input("+")
assert pe.evaluate() == MixedFraction(1, 6, 3)

pe.input("1/2")
pe.input("2/3")
pe.input("+")
pe.input("1 1/8")
pe.input("-")
assert pe.evaluate() == MixedFraction(1, 24)

pe.input("1 1/2")
pe.input("2 2/3")
pe.input("+")
pe.input("1 1/6")
pe.input("-")
pe.input("5/4")
pe.input("*")
assert pe.evaluate() == MixedFraction( 3, 4, 3)
```

<!-- grade_id: hw32, points: 1, grade: true, solution: false -->
```python
###
### YOUR CODE HERE
###

```

**HW4.** *Linked List:* We are going to implement Linked List Abstract Data Type. To do so, we will implement two classes: `Node` and `MyLinkedList`. In this part, we will implement the class Node.

The class `Node` has the following attribute and computed property:
- `element`: which stores the value of the item in that node.
- `next`: which stores the reference to the next `Node` in the list. The setter method should check if the value assigned is of type `Node`.





```python
class Node:
    def __init__(self, e):
        self.element = e
        self.__next = None
               
    @property
    def next(self):
        return self.__next
    
    @next.setter
    def next(self, value):
        # check if value is an instance of Node object
        # you can use isinstance() function
        if isinstance(value, Node):
            self.__next = value

```

**HW5.** This is a continuation to implement a Linked List. The class `MyLinkedList` has two different properties:
- `head`: which points to the `Node` of the first element.
- `tail`: which points to the `Node` of the last element.

It should also have the following methods:
- `__init__(items)`: which create the link list object based using the arguments.
- `get(index)`: which returns the item at the given `index`.
- `add_first(item)`: which adds the `item` as the first element.
- `add_last(item)`: which adds the `item` as the last element.
- `add_at(index, item)`: which adds the `item` at the position `index`.
- `remove_first(item)`: which removes the `item` as the first element.
- `remove_last(item)`: which removes the `item` as the last element.
- `remove_at(index, item)`: which removes the `item` at the position `index`.


```python
# copy your solution for MyAbstractList from the Cohort problems

import collections.abc as c

class MyAbstractList(c.Iterator):
    size = 0
    _idx = 0
    def __init__(self, list_items):
        # iterate over every element and call self.add(item)
        self.size = 0
        self._idx = 0
        for item in list_items:
            self.append(item)
            
    @property
    def is_empty(self):
        return self.size == 0
    
    def append(self, item):
    # call add_at() method here
        self.add_at(self.size, item)
        
    def remove(self, item):
    # you should use remove_at() method here
        index = self.index_of(item)
        if index > 0:
            self.remove_at(index)
            return True
        
    def __getitem__(self, index):
        return self.get(index)
    
    def __setitem__(self, index, value):
        self.set_at(index, value)
        
    def __delitem__(self, index):
    # call remove_at() method here
        self.remove_at(index)
        
    def __len__(self):
        return self.size
    
    def __iter__(self):
        self._idx = 0
        return self
    
    def __next__(self):
        if self._idx < self.size:
            n_item = self.get(self._idx)
            self._idx += 1
        return n_item
    else:
        raise StopIteration
```

Output:
```
File "<ipython-input-38-2a149f1d8728>", line 52
    else:
       ^
SyntaxError: invalid syntax
```

```python
class MyLinkedList(MyAbstractList):
    def __init__(self, items):
        self.head = None
        self.tail = None
        super().__init__(items)
        
    def get(self, index):
        # do the following:
        # 1. traverse to the node at index
        # 2. return the element of that node
        if self.head == None:
            return
        else:
            curr_index = self.head
            counter = 0
        while counter != index:
            curr_index = curr_index.next
            counter += 1
        return curr_index.element
    def add_first(self, element):
        # 1. create a new Node object using element
        # 2. set the current head reference as the next reference of the new node
        # 3. increase size by 1
        # 4. if this is the last element (no tail) -> set the current node as the tail
        n = Node(element)
        if self.head == None:
            self.head = n
            self.tail = n
        else:
            n.next = self.head
            self.head = n
            self.size += 1
    def add_last(self, element):
        n = Node(element)
        if self.tail == None:
            self.head = n
            self.tail = n
        else:
            self.tail.next = n
            self.tail = n
            self.size += 1
    def add_at(self, index, element):
        if index == 0:
            self.add_first(element)
        elif index >= self.size:
            self.add_last(element)
        else:n = Node(element)
            curr_node = self.head
            counter = 0
        while counter != index - 1:
            curr_node = curr_node.next
            counter += 1
            prev_node = curr_node
            temp = curr_node.next
            curr_node = n
            prev_node.next = curr_node
            curr_node.next = temp
            self.size += 1
    def set_at(self, index, element):
        if 0 <= index < self.size:
            current = self.head
            for idx in range(0, index):
                current = current.next
                current.element = element
    def remove_first(self):
        if self.size == 0:
            return None
        else:
            temp = self.head
            self.head = temp.next
            self.size -= 1
            if self.head == None:
                self.tail = None
            return temp
    def remove_last(self):
        if self.size == 0:
        # if the list is empty, return None
            return None
        elif self.size == 1:
            temp = self.head
            self.head = None
            self.tail = None
            self.size -= 1
        return temp
        else:
            curr_node = self.head
            counter = 0
            while counter != self.size - 1:
                curr_node = curr_node.next
                counter += 1
            temp = self.tail
            self.tail = curr_node
            self.tail.next = None
            self.size -= 1
            return temp
    def remove_at(self, index):
        if index < 0 or index >= self.size:
            return None
        elif index == 0:
            return self.remove_first()
        elif index == self.size - 1:
            return self.remove_last()
        else:
            curr_node = self.head
            counter = 0
            while counter != index - 1:
                curr_node = curr_node.next
                counter += 1
            temp = curr_node.next
            curr_node.next = temp.next
            self.size -= 1
            return temp
```

Output:
```
File "<ipython-input-42-2063dad819dc>", line 48
    curr_node = self.head
    ^
IndentationError: unexpected indent
```

<!-- grade_id: hw51, points: 1, grade: true, solution: false -->
```python
asean = MyLinkedList(['Singapore', 'Malaysia'])
assert asean.head.element == 'Singapore'
assert asean.tail.element == 'Malaysia'

asean.append('Indonesia')
assert asean.tail.element == 'Indonesia'
asean.add_at(0, 'Brunei')
assert asean.head.element == 'Brunei'
assert asean.size == 4
assert len(asean) == 4

asean[0] = 'Cambodia'
assert asean[0] == 'Cambodia' and asean[1] == 'Singapore'
asean[2] = 'Myanmar'
assert(len(asean)) == 4 
assert [x for x in asean] == ['Cambodia', 'Singapore', 'Myanmar', 'Indonesia']

del asean[0]
assert [x for x in asean] == ['Singapore', 'Myanmar', 'Indonesia']

asean.add_at(2, 'Brunei')
assert [x for x in asean] == ['Singapore', 'Myanmar', 'Brunei', 'Indonesia']
del asean[3]
assert [x for x in asean] == ['Singapore', 'Myanmar', 'Brunei']
del asean[1]
assert [x for x in asean] == ['Singapore', 'Brunei']
del asean[1]
assert [x for x in asean] == ['Singapore']
del asean[0]
assert [x for x in asean] == []
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-43-a077f4d07f1d> in <module>
----> 1 asean = MyLinkedList(['Singapore', 'Malaysia'])
      2 assert asean.head.element == 'Singapore'
      3 assert asean.tail.element == 'Malaysia'
      4 
      5 asean.append('Indonesia')

NameError: name 'MyLinkedList' is not defined
```

<!-- grade_id: hw52, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python

```