# Week 1 Problem Set
## Homework

**HW0.** Write a function `palindrome(s)`to check if the string `s` is a [Palindrome](https://en.wikipedia.org/wiki/Palindrome). Use iteration for this solution.


```python
def palindrome(s):
    if s == s[::-1]:
        return True
    else:
        return False
```

<!-- grade_id: hw01, points: 1, grade: true, solution: false -->
```python
assert not palindrome("moon") 
assert palindrome("noon") 
assert palindrome("a a") 
assert palindrome("ada") 
assert not palindrome("ad a") 
```

**HW1.** Create a function to implement **Bubble Sort Algorithm version 2**. You can check the pseudocode in [Bubble Sort and Insertion Sort](https://github.com/Data-Driven-World/d2w_notes/blob/master/BubbleSort_InsertionSort.ipynb). 

The function should ...

- **modify/mutate** an existing array into a sorted one
- **return** the number of comparisons made

Hint: To count the number of comparisons made you can simply increment an integer variable (counter) right before the `if` statement.

```python
def bubble_sort(array):
    n=len(array)
    count = 0
    swap=True
    while swap == True:
        swap = False
        for i in range(1,n):
            count+=1
            print(array)
            if array[i] < array[i-1]:
                array[i], array[i-1]= array[i-1], array[i]
                swap=True
    return count
```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
array = [10,3,8,47,1,0,-39,8,4,7,6,-5]
count = bubble_sort(array)
print(array, count)
assert array == [-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47]
assert count == 121
```

Output:
```
[10, 3, 8, 47, 1, 0, -39, 8, 4, 7, 6, -5]
[3, 10, 8, 47, 1, 0, -39, 8, 4, 7, 6, -5]
[3, 8, 10, 47, 1, 0, -39, 8, 4, 7, 6, -5]
[3, 8, 10, 47, 1, 0, -39, 8, 4, 7, 6, -5]
[3, 8, 10, 1, 47, 0, -39, 8, 4, 7, 6, -5]
[3, 8, 10, 1, 0, 47, -39, 8, 4, 7, 6, -5]
[3, 8, 10, 1, 0, -39, 47, 8, 4, 7, 6, -5]
[3, 8, 10, 1, 0, -39, 8, 47, 4, 7, 6, -5]
[3, 8, 10, 1, 0, -39, 8, 4, 47, 7, 6, -5]
[3, 8, 10, 1, 0, -39, 8, 4, 7, 47, 6, -5]
[3, 8, 10, 1, 0, -39, 8, 4, 7, 6, 47, -5]
[3, 8, 10, 1, 0, -39, 8, 4, 7, 6, -5, 47]
[3, 8, 10, 1, 0, -39, 8, 4, 7, 6, -5, 47]
[3, 8, 10, 1, 0, -39, 8, 4, 7, 6, -5, 47]
[3, 8, 1, 10, 0, -39, 8, 4, 7, 6, -5, 47]
[3, 8, 1, 0, 10, -39, 8, 4, 7, 6, -5, 47]
[3, 8, 1, 0, -39, 10, 8, 4, 7, 6, -5, 47]
[3, 8, 1, 0, -39, 8, 10, 4, 7, 6, -5, 47]
[3, 8, 1, 0, -39, 8, 4, 10, 7, 6, -5, 47]
[3, 8, 1, 0, -39, 8, 4, 7, 10, 6, -5, 47]
[3, 8, 1, 0, -39, 8, 4, 7, 6, 10, -5, 47]
[3, 8, 1, 0, -39, 8, 4, 7, 6, -5, 10, 47]
[3, 8, 1, 0, -39, 8, 4, 7, 6, -5, 10, 47]
[3, 8, 1, 0, -39, 8, 4, 7, 6, -5, 
```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** The solution in HW1 can be improved!

The n-th pass places the n-th largest element into its final place (i.e. after 1 cycle, the (1st) largest element will be in its correct position).

Hence, we can reduce subsequent iterations by only considering the other (unsorted) `n-1` elements in the array. 

Implement the algorithm of **Bubble Sort version 4** under [Bubble Sort and Insertion Sort](https://github.com/Data-Driven-World/d2w_notes/blob/master/BubbleSort_InsertionSort.ipynb). 


```python
def bubble_sort(array):
    n=len(array)
    count = 0
    swap=True
    ending = n
    while swap == True:
        swap = False
        for i in range(1,ending):
            count+=1
            print(array)
            if array[i] < array[i-1]:
                array[i], array[i-1]= array[i-1], array[i]
                swap=True
        ending -= 1
    return count
```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
array = [10,3,8,47,1,0,-39,8,4,7,6,-5]
count = bubble_sort(array)
print(array, count)
assert array == [-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47]
assert count == 66
```

Output:
```
[10, 3, 8, 47, 1, 0, -39, 8, 4, 7, 6, -5]
[3, 10, 8, 47, 1, 0, -39, 8, 4, 7, 6, -5]
[3, 8, 10, 47, 1, 0, -39, 8, 4, 7, 6, -5]
[3, 8, 10, 47, 1, 0, -39, 8, 4, 7, 6, -5]
[3, 8, 10, 1, 47, 0, -39, 8, 4, 7, 6, -5]
[3, 8, 10, 1, 0, 47, -39, 8, 4, 7, 6, -5]
[3, 8, 10, 1, 0, -39, 47, 8, 4, 7, 6, -5]
[3, 8, 10, 1, 0, -39, 8, 47, 4, 7, 6, -5]
[3, 8, 10, 1, 0, -39, 8, 4, 47, 7, 6, -5]
[3, 8, 10, 1, 0, -39, 8, 4, 7, 47, 6, -5]
[3, 8, 10, 1, 0, -39, 8, 4, 7, 6, 47, -5]
[3, 8, 10, 1, 0, -39, 8, 4, 7, 6, -5, 47]
[3, 8, 10, 1, 0, -39, 8, 4, 7, 6, -5, 47]
[3, 8, 10, 1, 0, -39, 8, 4, 7, 6, -5, 47]
[3, 8, 1, 10, 0, -39, 8, 4, 7, 6, -5, 47]
[3, 8, 1, 0, 10, -39, 8, 4, 7, 6, -5, 47]
[3, 8, 1, 0, -39, 10, 8, 4, 7, 6, -5, 47]
[3, 8, 1, 0, -39, 8, 10, 4, 7, 6, -5, 47]
[3, 8, 1, 0, -39, 8, 4, 10, 7, 6, -5, 47]
[3, 8, 1, 0, -39, 8, 4, 7, 10, 6, -5, 47]
[3, 8, 1, 0, -39, 8, 4, 7, 6, 10, -5, 47]
[3, 8, 1, 0, -39, 8, 4, 7, 6, -5, 10, 47]
[3, 8, 1, 0, -39, 8, 4, 7, 6, -5, 10, 47]
[3, 1, 8, 0, -39, 8, 4, 7, 6, -5, 
```

<!-- grade_id: hw22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```