# Week 6 Problem Set

## Cohort Sessions

**CS1.** Create a class called `Fraction` to represent a simple fraction. The class has two properties:
- `num`: which represents a numerator and of the type Integer.
- `den`: which represents a denominator and of the type Integer. Denominator should not be a zero. If a zero is assigned, you need to replace it with a `1`. 

The class should have the following method:
- `__init__(num, den)`: to initialize the numerator and the denominator. You should check if the denominator is zero. If it is you should assign `1` as the denominator instead. 
- `__str__()`:  for the object instance to be convertable to String.  You need to return a string in a format of `num/den`


```python
class Fraction:
    def __init__(self, num, den):
        ###BEGIN SOLUTION
        self.num = num # set attributes via the SETTER
        self.den = den
        ###END SOLUTION
        pass
    
    @property
    def num(self):
        ###BEGIN SOLUTION
        return self._num
        ###END SOLUTION
        pass
    
    @num.setter
    def num(self, val):
        ###BEGIN SOLUTION 
        self._num = int(val) # this creates attribute _num IF it does not exist already, else this modifies existing attr _num
        ###END SOLUTION
        pass
    
    @property
    def den(self):
        ###BEGIN SOLUTION
        return self._den
        ###END SOLUTION
        pass
    
    @den.setter
    def den(self, val):
        ###BEGIN SOLUTION
        if (int(val) == 0):
            self._den = 1
        else:
            self._den = int(val)
        ###END SOLUTION
        pass
    
    def __str__(self):
        ###BEGIN SOLUTION
        # 3/5
        return f"{self.num:d}/{self.den:d}"
        ###END SOLUTION
        pass
    
```

<!-- grade_id: cs11, points: 1, grade: true, solution: false -->
```python
f0 = Fraction(0, 1)
assert f0.num == 0
assert f0.den == 1
assert str(f0) == "0/1"

f1 = Fraction(1, 2)
assert f1.num == 1
assert f1.den == 2
assert str(f1) == "1/2"

f1.num = 3
f1.den = 4
assert str(f1) == "3/4"
```

<!-- grade_id: cs12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS2.** Extend the class `Fraction` to support the following operator: `+` and `==`. To do this, you need to overload the following operator:
- `__add__(self, other)`
- `__eq__(self, other)`

You may want to write a method to simplify a fraction:
- `simplify()`: which simplify a fraction to its lowest terms. To simplify a fraction divide both the numerator and the denominator with the greatest common divisor of the the two. This method should return a new `Fraction` object.


```python
def gcd(a, b):
    if b == 0:
        return a
    else:
        return gcd(b, a % b)


class Fraction:
    # copy all the other methods from the previous exercies
    
    ###BEGIN SOLUTION
    def __init__(self, num, den):
        ###BEGIN SOLUTION
        self.num = num # set attributes via the SETTER
        self.den = den
        ###END SOLUTION
        pass
    
    @property
    def num(self):
        ###BEGIN SOLUTION
        return self._num
        ###END SOLUTION
        pass
    
    @num.setter
    def num(self, val):
        ###BEGIN SOLUTION 
        self._num = int(val) # this creates attribute _num IF it does not exist already, else this modifies existing attr _num
        ###END SOLUTION
        pass
    
    @property
    def den(self):
        ###BEGIN SOLUTION
        return self._den
        ###END SOLUTION
        pass
    
    @den.setter
    def den(self, val):
        ###BEGIN SOLUTION
        if (int(val) == 0):
            self._den = 1
        else:
            self._den = int(val)
        ###END SOLUTION
        pass
    
    def __str__(self):
        ###BEGIN SOLUTION
        # 3/5
        return f"{self.num:d}/{self.den:d}"
        ###END SOLUTION
        pass
    ###END SOLUTION
    
    # @return: a new instance of simplified Fraction
    def simplify(self):
        ###BEGIN SOLUTION
        gcd_value = gcd(self.num, self.den)
        simplified_numerator = self.num / gcd_value
        simplified_denominator = self.den / gcd_value
        return Fraction(simplified_numerator, simplified_denominator)
        ###END SOLUTION
        pass
    
    # @return: a new instance of SIMPLIFIED Fraction (self + other)
    def __add__(self, other):
        ###BEGIN SOLUTION
        new_num = self.num * other.den + self.den * other.num
        new_den = self.den * other.den 
        result = Fraction(new_num, new_den).simplify()
        return result 
        ###END SOLUTION
        pass
        
    
    def __eq__(self, other):
        print("eq func is called")
        ###BEGIN SOLUTION
        f1 = self.simplify()
        f2 = other.simplify()
        return f1.num == f2.num and f1.den == f2.den
        ###END SOLUTION
    

```

<!-- grade_id: cs21, points: 1, grade: true, solution: false -->
```python
f1 = Fraction(1, 2)
f2 = Fraction(2, 3)
f3 = f1 + f2 

assert str(f3) == "7/6"

f4 = Fraction(3, 5)
f5 = Fraction(1, 3)
f6 = f4 + f5 

assert str(f6) == "14/15"

f1 = Fraction(1, 2)
f2 = Fraction(2, 4)
assert f1 == f2

f3 = Fraction(2, 3)
f4 = Fraction(2, 4)
assert f3 != f4
```

Output:
```
eq func is called
eq func is called
```

<!-- grade_id: cs22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS3.** *Inheritance:* Create a class called `MixedFraction` as a subclass of `Fraction`. A mixed fraction is a fraction that comprises of a whole number, a numerator and a denominator, e.g. `1 2/3` which is the same as `5/3`. The class has the following way of initializing its properties:
- `__init__(top, bot, whole)`: which takes in three Integers, the whole number, the numerator, and the denominator, e.g. `whole=1`, `top=2`, `bot=3`. The argument `whole` by default is `0`.  You can also specify `top` to be greater than `bot`. 

The class only has two properties:
- `num`: which is the numerator and can be greater than denominator.
- `den`: which is the denominator and must be a non-zero number.

The class should also have the following methods:
- `get_three_numbers()`: which is used to calculate the whole number, numerator and the denominator from a given numerator and denominator. The stored properties are `num` and `den` as in `Fraction` class. This function returns three Integers as a tuple, i.e. `(top, bot, whole)`.

The class should also override the `__str__()` method in this manner:
- `num/dem` if the numerator is smaller than the denominator. For example, `2/3`. 
- `whole top/bot` if the numerator is greater than the denominator. For example, `1 2/3`.

```python
class MixedFraction(Fraction):
    def __init__(self, top, bot, whole=0):
        ###BEGIN SOLUTION
        # 1 1/2 => 3/2
        num = top + bot * whole
        den = bot
        ###END SOLUTION
        super().__init__(num, den)

    def get_three_numbers(self):
        whole = None
        top = None
        bot = None
        ###BEGIN SOLUTION
        whole = self.num // self.den # int division
        top = self.num % self.den # remainder
        bot = self.den
        ###END SOLUTION
        return (top, bot, whole)

    def __str__(self):
        ###BEGIN SOLUTION
        (top, bot, whole) = self.get_three_numbers()
        if whole == 0:
            return super().__str__()
        return f"{whole:d} {top:d}/{bot:d}"
        ###END SOLUTION
        pass
```

<!-- grade_id: cs31, points: 1, grade: true, solution: false -->
```python
mf1 = MixedFraction(5, 3)
assert mf1.num == 5 and mf1.den == 3
assert mf1.get_three_numbers() == (2, 3, 1)
mf2 = MixedFraction(2, 3, 1)
assert mf2.num == 5 and mf2.den == 3

result = mf1 + mf2
assert result.num == 10 and result.den == 3

assert mf1 == mf2
```

Output:
```
eq func is called
```

<!-- grade_id: cs32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS4.** *Inheritance:* Create a class `Deque` as a subclass of `Queue`. Use the double-stack implementation of `Queue` in this problem. `Deque` has the following methods:
- `add_front(item)`: which add an item to the front of the queue. 
- `remove_rear()`: which pops out an item from the rear of the queue. 
- `add_rear(item)`: which add an item from rear of the queue. This is the same as enqueue a normal queue.
- `remove_front()`: which pops out an item from the front of the queue. This is the same as dequeue method in a normal queue.

See the notation for the front and rear of a Queue.

![](https://www.dropbox.com/s/xexwpbhljopmg37/queue.png?raw=1)


```python
class Stack:
    ###BEGIN SOLUTION
    def __init__(self):
        self.items = []

    @property
    def is_empty(self):
        return self.items == []

    @property
    def size(self):
        return len(self.items)

    def push(self, item):
        self.items.append(item)

    def pop(self):
        if len(self.items) >= 1:
            return self.items.pop()

    def peek(self):
        if len(self.items) >= 1:
            return self.items[-1]
    ###END SOLUTION
```

```python
class Queue:
    def __init__(self):
        self.left_stack = Stack()
        self.right_stack = Stack()
    
    ###BEGIN SOLUTION
    @property
    def is_empty(self):
        return self.left_stack.is_empty and self.right_stack.is_empty
    
    @property
    def size(self):
        return self.left_stack.size + self.right_stack.size
    
    def enqueue(self, item):
        self.right_stack.push(item)
        
    def right_to_left(self):
        while not self.right_stack.is_empty:
            self.left_stack.push(self.right_stack.pop())
    
    def dequeue(self):
        if self.left_stack.is_empty:
            self.right_to_left()
        return self.left_stack.pop()
    
    def peek(self):
        if self.left_stack.is_empty:
            if not self.right_stack.is_empty:
                self.right_to_left()
                return self.left_stack.peek()
        else:
            return self.left_stack.peek()
    ###END SOLUTION
```

```python
class Dequeue(Queue):
  
    # cut queue
    # add items to the front of the queue -> the top of the left stack
    def add_front(self, item):
        ###BEGIN SOLUTION
        self.left_stack.push(item)
        ###END SOLUTION
        pass
      
    # @return: item in the Queue
    def remove_front(self):
        ###BEGIN SOLUTION
        return self.dequeue()
        ###END SOLUTION
        pass
    
    # @return: None
    def add_rear(self, item):
        ###BEGIN SOLUTION
        self.enqueue(item)
        ###END SOLUTION
        pass
    
    def left_to_right(self):
        ###BEGIN SOLUTION
        while not self.left_stack.is_empty:
            self.right_stack.push(self.left_stack.pop())
        ###END SOLUTION
        pass
    
    
    def remove_rear(self):
        ###BEGIN SOLUTION
        if self.right_stack.is_empty:
            self.left_to_right()
        return self.right_stack.pop()
        ###END SOLUTION
        pass
    
    def peek_front(self):
        ###BEGIN SOLUTION
        return self.peek()
        ###END SOLUTION
        pass
    
    def peek_rear(self):
        ###BEGIN SOLUTION
        if self.right_stack.is_empty:
            self.left_to_right()
        return self.right_stack.peek()
        ###END SOLUTION
        pass
```

<!-- grade_id: cs41, points: 1, grade: true, solution: false -->
```python
q1 = Dequeue()
q1.add_front(1)
q1.add_front(2)
q1.add_front(3)
assert q1.remove_rear()==1
assert q1.remove_rear()==2
assert q1.remove_rear()==3

q1.add_rear(3)
q1.add_rear(2)
q1.add_rear(1)
assert q1.remove_front()==3
assert q1.remove_front()==2
assert q1.remove_front()==1

q1.add_front(1)
q1.add_rear(2)
q1.add_front(3)
q1.add_rear(4)
q1.add_front(5)
q1.add_rear(6)
assert q1.remove_rear()==6
assert q1.remove_rear()==4
assert q1.remove_rear()==2
assert q1.remove_rear()==1
assert q1.remove_rear()==3
assert q1.remove_front()==5
```

<!-- grade_id: cs42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS5-Prelude.** *ArrayFixedSize class (Given):* Write a class called `ArrayFixedSize` that simulate a fixed size array. This class should inherint from `collections.abc.Iterable` base class. A fixed size array is like a list which size cannot change once it is set. The size and its data type is specified during object instantiation. Use Numpy's array for its internal data storage. At the start you can use `np.empty(size)` to create an uninitalized empty array. The class should have the following methods:

- `__getitem__(index)`: which returns the element at a given `index` using the bracket operator, i.e. `array[index]`.
- `__setitem__(index, value)`: which set the item at a given `index` with a particular `value` using the bracket and assignment operators, i.e. `array[index] = value`.
- `__iter__()`: which returns the iterable object so that it can be used in a for loop and other iterable object operators.
- `__str__()`: which returns the string object representation of the object. It should displays as follows: `[el1, el2, el3, ...]`.
- `__len__()`: which returns the number of items in the array when `len()` function is called upon this object.

```python
import  collections.abc as c
import numpy as np

class ArrayFixedSize(c.Iterable):
    
    def __init__(self, size, dtype=int):
        self.__data = np.empty(size)
        self.__data = self.__data.astype(dtype)
        
    def __getitem__(self, index):
        print("get item is called")
        return self.__data[index]
    
    def __setitem__(self, index, value):
        print("set item is called")
        self.__data[index] = value      
        
    def __iter__(self):
        return iter(self.__data)
    
    def __len__(self):
        return len(self.__data)
        
    def __str__(self):
        out = "["
        for item in self:
            out += f"{item:}, "
        if self.__data != []:
            return out[:-2] + "]"
        else:
            return "[]"

x = ArrayFixedSize(5)
print(x)
# try to set
x[3] = 123 # a special way to call __setitem__
print(x[2]) # a special way to call __getitem__
len(x) # a special way to call __len__
for i in x: # a special way to call __iter__
    print(i) 
print(x)
```

Output:
```
[0, 0, 0, 0, 0]
set item is called
get item is called
0
0
0
0
123
0
[0, 0, 0, 123, 0]
/usr/lib/python3.7/site-packages/ipykernel_launcher.py:28: DeprecationWarning: elementwise comparison failed; this will raise an error in the future.
```

**CS5.** Implement a class called `MyAbstractList` which is a subclass of Python's `collections.abc.Iterator` class. This class is meant to be an abstract class for two kinds of List data structures, i.e. `MyArrayList` and `MyLinkedList`. In this exercise, you will implement `MyAbstractList`.

This class has the following attribute and computed property property:
- `size`: which gives you the size of the list in integer.
- `is_empty`: which returns either `True` or `False` depending whether the list is empty or not.

Implement also several other methods:
- `__init__(self, list_of_items)`: which initializes the list by adding the items in list argument using the `add(item)` method.
- `__getitem__(index)`: which returns the element at a given `index` using the bracket operator, i.e. `array[index]`. This method should call the child's method `get(index)`.
- `__setitem__(index, value)`: which set the item at a given `index` with a particular `value` using the bracket and assignment operators, i.e. `array[index] = value`. This method should call the child's method `set_at(index, item)`.
- `__delitem__(index)`: which removes the item at a given `index` using the `del` operator, i.e. `del array[index]`. This method should call the child's method `remove_at(index)`.
- `append(item)`: which adds an item at the end of the list. This method should call `add_at(index, item)` in the child class which adds the item at the specified index. 
- `remove(item)`: which removes the first occurence of the item in the list. This method should call `remove(index)` in the child class which removes the item at the specified index.
- `__next__()`: which returns the next element in the iterator. This method should call `get(index)` in the child class which returns the item at the specified index. If there is no more element, it should `raise StopIteration`.

Create any other attributes that you think is necessary.





```python
import collections.abc as c

# an abstract class by defn should not contain implementation
# however this cohort exercise ....
class MyAbstractList(c.Iterator):
    
    def __init__(self, list_items):
        # iterate over every element and call self.add(item)
        self.size = 0
        self._idx = 0
        # where is the add method?
        # it's not add, it's append
        # typo
        for item in list_items:
            self.append(item)
        ###END SOLUTION
    
    # at least specify them here
    # so the child class know that
    # we need to have these methods
    def add_at(self, idx, item):
        pass
    def index_of(self, item):
        pass
    def remove_at(self, idx, item):
        pass
    def get(self, idx):
        pass
    def set_at(self, idx, item):
        pass

    @property
    def is_empty(self):
        ###BEGIN SOLUTION
        return self.size == 0
        ###END SOLUTION
        pass
    
    def append(self, item):
        # call add_at() method here
        # where is add_at()?
        # add_at()  is at the CHILD class
        # how do i know that my child class will implement add_at()??
        # i dont... 
        self.add_at(self.size, item)
        # if self == MyAbstractList then, we don't have add_at
        # so self must be the CHILD of MyAbstractList, 
        # that implements add_at() function 
        # NO!!!!!
        ###BEGIN SOLUTION
        ###END SOLUTION
        
    def remove(self, item):
        # you should use remove_at() method here
        # must find the index of item in the list
        idx_item = self.index_of(item) # child class method
        if idx_item >= 0:
            self.remove_at(idx_item) # child class method
            return True
        else:
            return False
        ###BEGIN SOLUTION

        ###END SOLUTION
        
    def __getitem__(self, index):
        ###BEGIN SOLUTION
        return self.get(index)
        # self would be an instance of MyAbstractList CHILD instance
        ###END SOLUTION
        pass
    
    def __setitem__(self, index, value):
        # call set_at(index, value) method here
        ###BEGIN SOLUTION
        self.set_at(index, value)
        ###END SOLUTION
        pass
        
    def __delitem__(self, index):
        # call remove_at() method here
        ###BEGIN SOLUTION
        self.remove_at(index)
        ###END SOLUTION
        pass
    
    def __len__(self):
        ###BEGIN SOLUTION
        return self.size
        ###END SOLUTION
        pass
        
    def __iter__(self):
        self._idx = 0
        return self
        
    def __next__(self):
        if self._idx < self.size:
            n_item = self.get(self._idx)
            self._idx += 1
            return n_item
        else:
            raise StopIteration
```

```python
# creating a class PythonList inheriting from MyAbstractList
# this is just for testing the MyAbstractList class

class PythonList(MyAbstractList):
    data = []
    
    def __init__(self):
        self.data = []
        super().__init__(list(range(10)))
    
    def add_at(self, index, item):
        self.data.insert(index, item)
        self.size += 1
        
    def set_at(self, index, item):
        self.data[index] = item
        
    def remove_at(self, index):
        self.data.pop(index)
        self.size -= 1
        
    def get(self, index):
        if 0 <= index < self.size:
            return self.data[index]
        else:
            raise IndexError()
            
    def index_of(self, item):
        try:
            idx = self.data.index(item)
            return idx
        except:
            return -1

```

<!-- grade_id: cs51, points: 1, grade: true, solution: false -->
```python
f = PythonList()

# Testing init
assert f.data == list(range(10))

# Testing size property
assert f.size == 10

# Testing is_empty property
assert not f.is_empty

# Testing add method
f.append(10)
assert f.data == list(range(11))

# Testing remove method
f.remove(5)
assert f.data == [0, 1, 2, 3, 4, 6, 7, 8, 9, 10]
f.add_at(5, 5)
f.append(5)
f.remove(5)
assert f.data == [0, 1, 2, 3, 4, 6, 7, 8, 9, 10, 5]
assert not f.remove(11)

# Testing getitem method
assert [f[i] for i in range(len(f))] == [0, 1, 2, 3, 4, 6, 7, 8, 9, 10, 5]
f[0] = -1
assert [f[i] for i in range(len(f))] == [-1, 1, 2, 3, 4, 6, 7, 8, 9, 10, 5]

# Testing delitem
del f[0]
assert f[0] == 1

# Testing __next__
assert [x for x in f] == [ 1, 2, 3, 4, 6, 7, 8, 9, 10, 5]
# Testing __iter__
assert [x for x in f] == [ 1, 2, 3, 4, 6, 7, 8, 9, 10, 5]
```

**CS6.** Implement the class `MyArrayList` which is a subclass of `MyAbstractList`. `MyArrayList` should implement the `ArrayFixedSize` class instead of Python's list as in `PythonList` class above. You need to implement the following methods:
- `__init__(self, items)`: where items are used to initialize the array.
- `add_at(self, index, item)`: which add the `item` at a particular `index`. The algorithm for this method is as follows:
    - calls `ensure_capacity()` which doubles the size of the array into a new array and copies the data.
    - shift all the items by one for every position after the `index`.
    - add the `item` at `index`.
    - add the `size` attribute by 1.
- `set_at(self, index, value)`: this method replaces the item at `index` with a new value. 
- `remove_at(self, index)`: which removes the item at `index`.
- `get(self, index)`: which returns the item at `index` if the index is valid, otherwise, it raise `IndexError()`.
- `index_of(self, item)`: which return the index of the given `item`, otherwise, it returns `-1`.
- `clear(self)`: which set the data back to zero with the size of the initial capacity.

Since we are using `ArrayFixedSize` as the internal data type, take a look at the following hints:
- We will create the array data in blocks of 16. For example, in the beginning, even if you have data less 16, we will create 16 data blocks. We will only initialize the value in that block to what we need.
- When we add data into the list, we need to ensure we have enough capacity. So when the initialized data has reach the end of 16 blocks, we need to create a new block with double the size. You then need to copy the data from the old block to the new block. 

```python
import numpy as np

# trying to mock FixedArray implementation
class MyArrayList(MyAbstractList):
    INITIAL_CAPACITY = 16
    
    def __init__(self, items, dtype=int):
        self.data = ArrayFixedSize(MyArrayList.INITIAL_CAPACITY, dtype)
        # iterate over every items and call add(item)
        ###BEGIN SOLUTION
        super().__init__(items)
        ###END SOLUTION
    
    def add_at(self, index, item):
        self.ensure_capacity()
        # do the following:
        # 1. copy data by shifting it to the right from index position to the end
        # 2. set item at index
        # 3. add size by 1
        ###BEGIN SOLUTION
        for idx in range(self.size-1, index-1, -1):
            self.data[idx+1] = self.data[idx] 
            # shift elements to the right
            # copy to the right
        # insert the new item at the index location
        self.data[index] = item # calls the __setitem__ of ArrayFixedSize
        # update the fixed array size
        self.size += 1

        ###END SOLUTION
        
    def set_at(self, index, value):
        ###BEGIN SOLUTION
        self.data[index] = value
        ###END SOLUTION
        pass
        
    def remove_at(self, index):
        if 0 <= index < self.size:
            # do the following
            # 1. get the element at index
            # 2. copy the data by shifting it to the left from index to the end
            # 3. return the element at index
            ###BEGIN SOLUTION
            element = self.data[index]
            for idx in range(index, self.size-1):
                # shift everything to the left
                self.data[idx] = self.data[idx+1]
            # update size to -1
            self.size -= 1
            return element
            ###END SOLUTION
        else:
            raise IndexError()
        
    def get(self, index):
        if 0 <= index < self.size:
            ###BEGIN SOLUTION
            # rely on __getitem__ method of ArrayFixedSize
            return self.data[index]
            ###END SOLUTION
        else:
            raise IndexError()
            
    def index_of(self, item):
        # given an item (element), find the index in the list
        # if doesnt exist, return -1 
        for idx in range(0, self.size):
            if self.data[idx] == item:
                return idx 
        return -1
        
    def ensure_capacity(self):
        if self.size >= len(self.data):
            new_data = ArrayFixedSize(self.size * 2 + 1)
            self.copy(self.data, 0, new_data, 0)
            self.data = new_data
            
    def copy(self, source, idx_s, dest, idx_d):
        for idx in range(idx_s,len(source)):
            offset = idx - idx_s
            dest[idx_d + offset] = source[idx]
            
    def clear(self):
        self.data = ArrayFixedSize(MyArrayList.INITIAL_CAPACITY)
        self.size = 0
        
    def __str__(self):
        out = "["
        for idx in range(self.size):
            out += f"{self.get(idx):}, "
        return out[:-2] + "]"
```

<!-- grade_id: cs61, points: 1, grade: true, solution: false -->
```python
a = MyArrayList([1,2,3])
assert [x for x in a] == [1,2,3]
assert a.size == 3
assert not a.is_empty
a.append(4)
assert a.size == 4
assert [x for x in a] == [1,2,3,4]
assert [a[i] for i in range(len(a))] == [1,2,3,4]
a[0] = -1
assert [x for x in a] == [-1, 2,3,4]
del a[0]
assert [x for x in a] == [2,3,4]
```

Output:
```
set item is called
set item is called
set item is called
get item is called
get item is called
get item is called
set item is called
get item is called
get item is called
get item is called
get item is called
get item is called
get item is called
get item is called
get item is called
set item is called
get item is called
get item is called
get item is called
get item is called
get item is called
get item is called
set item is called
get item is called
set item is called
get item is called
set item is called
get item is called
get item is called
get item is called
```

<!-- grade_id: cs62, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```