# Week 3 Problem Set

## Homeworks

**HW1.** *Fibonaci:* Write a function to find the Fibonacci number given an index. These are the example of a the first few numbers in Fibonacci series.

| Indices    | 0 | 1 | 2 | 3 | 4 | 5 | 6 | 7  | 8  | 9  |
|------------|---|---|---|---|---|---|---|----|----|----|
| The series | 0 | 1 | 1 | 2 | 3 | 5 | 8 | 13 | 21 | 34 |

We can write that the Fibonacci number at index *i* is given by:

$fibo(i) = fibo(i-1) + fibo(i-2)$

User recursion for your implementation.

```python
def fibonacci(index):
    if index-2 >=0:
        f_index = fibonacci(index-1) + fibonacci(index-2)
        return f_index
    elif index == 1:
        return 1
    else:
        return 0


```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
assert fibonacci(0) == 0
assert fibonacci(1) == 1
assert fibonacci(3) == 2
assert fibonacci(7) == 13
assert fibonacci(9) == 34
```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** *Max-Heapify:* Recall the algorithm for max-heapify in restoring the *max-heap property* of a binary heap tree. In the previous implementation, we used iteration. Implement the same algorithm using recursion. 

```python
def left_of(index):
    leftchild = int((index*2)+1)
    return leftchild

def right_of(index):
    rightchild = int((index*2)+2)
    return rightchild

def max_child(array, index, heap_size):
    if right_of(index) >= heap_size:
        return left_of(index)
    else:
        if array[right_of(index)] >= array[left_of(index)]:
            return right_of(index)
        else:
            return left_of(index)
        
def max_heapify_recursive(array, current_i , heap_size):
    if left_of(current_i) < heap_size:#size=10#
        max_child_i = max_child(array, current_i, heap_size)
        if array[max_child_i] > array[current_i]:
            array[max_child_i], array[current_i] = array[current_i], array[max_child_i]
        current_i = max_child_i
        max_heapify_recursive(array, current_i , heap_size)
        

def max_heapify(array, index, heap_size):
#     current_i = ((heap_size-2)//2)
    current_i = index #1
    max_heapify_recursive(array, current_i , heap_size)
        

    
        
result = [16, 4, 10, 14, 7, 9, 3, 2, 8, 1]
max_heapify(result, 1, len(result))
print(result)
assert result == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
result = [4, 1, 10, 14, 16, 9, 3, 2, 8, 7]
max_heapify(result, 1, len(result))
print(result)
assert result == [4, 16, 10, 14, 7, 9, 3, 2, 8, 1]     
```

Output:
```
[16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
[4, 16, 10, 14, 7, 9, 3, 2, 8, 1]
```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
# result = [16, 4, 10, 14, 7, 9, 3, 2, 8, 1]
max_heapify(result, 1, len(result))
print(result)
assert result == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
result = [4, 1, 10, 14, 16, 9, 3, 2, 8, 7]
max_heapify(result, 1, len(result))
print(result)
assert result == [4, 16, 10, 14, 7, 9, 3, 2, 8, 1]
```

Output:
```
[4, 16, 10, 14, 7, 9, 3, 2, 8, 1]
AssertionError                            Traceback (most recent call last)
<ipython-input-3-bdc31353cc02> in <module>
      2 max_heapify(result, 1, len(result))
      3 print(result)
----> 4 assert result == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
      5 result = [4, 1, 10, 14, 16, 9, 3, 2, 8, 7]
      6 max_heapify(result, 1, len(result))

AssertionError:
```

<!-- grade_id: hw22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW3.** *String Permutation:* Write two functions to return an array of all the permutations of a string. For example, if the input is `abc`, the output should be

```
output = ["abc", "acb", "bac", "bca", "cab", "cba"]
```

The first function only has one argument which is the input string, i.e. `permutate(s)`. The second function is the recursive function with two arguments String 1 (`str1`) and and String 2 (`str2`). The first function calls the second method `permutate("", s)` at the beginning. The second function should use a loop to move a character from `str2` to `str1` and recursively invokes it with a new `str1` and `str2`. 

```python

def permutate(s):
    lst=[]
    if len(s) == 0:
        lst.append(s)
    else:
        recursive("",s,lst)
        return lst
        
def recursive(str1,str2,lst):
    if str2 == "":
#         print(str1,"it works")
        lst.append(str(str1))
        return lst
    for i in range(0, len(str2)):
#         print(f"str1..{str1} str2..{str2}")
        recursive(str1+str2[i],str2[:i]+str2[1+i:],lst)

permutate("abc")
```

Output:
```
['abc', 'acb', 'bac', 'bca', 'cab', 'cba']
```

<!-- grade_id: hw31, points: 1, grade: true, solution: false -->
```python
assert set(permutate("abc"))==set(["abc", "acb", "bac", "bca", "cab", "cba"])
assert set(permutate("abcd"))==set(["abcd", "abdc", "acbd", "acdb", "adbc", "adcb", "bacd", "badc", "bcad", "bcda", "bdac", "bdca", "cabd", "cadb", "cbad", "cbda", "cdab", "cdba", "dabc", "dacb", "dbac", "dbca", "dcab", "dcba"])
```

<!-- grade_id: hw32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW4.** *GCD:* Implement the recursive implementation of Euclid's algorithm for Greatest Common Divisor (GCD).

Hint: [Wikipedia: Greatest Common Divisor](https://en.wikipedia.org/wiki/Greatest_common_divisor#Euclid's_algorithm)

```python
def gcd(a, b):
    if a == 0 or b == 0:
        return a+b
    elif a > b:
        return gcd(b,a%b)
    elif  a < b:
        return gcd(a,b%a)

# print(gcd(0,0))
#print(gcd(2,0))
# print(gcd(0,2))
# print(gcd(2,3))
print(gcd(4, 28))
gcd(1024, 526)
print(gcd(1024, 526))
```

Output:
```
4
2
```

<!-- grade_id: hw41, points: 1, grade: true, solution: false -->
```python
assert gcd(0,0) == 0
assert gcd(2,0) == 2
assert gcd(0,2) == 2
assert gcd(2,3) == 1
assert gcd(4, 28) == 4
assert gcd(1024, 526) == 2
```

<!-- grade_id: hw42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python

```